```python
import matplotlib.pyplot as plt

%matplotlib inline
%precision 16
import numpy

```

# Homework 0:  A Test

This notebook contains two short (and hopefully simple) problems to test the homework collection and grading system for APMA4300.  

### General Instructions:  

Just fill in the cells marked `YOUR ANSWER HERE` (Markdown) or `# YOUR CODE HERE` (python).  

**Please!** do not copy any of the existing cells as that can break the autograder.  If necessary, you can always _add_ additional cells but do not copy cells...particular ones that will be graded.

## Question 1  

In the cell below, write out in LaTeX, the quadratic formula for the solution to 

$$
    x^2 + \alpha x + \beta = 0
$$

<!-- grade_id: cell1448053803160613, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

Write a python function that returns the two roots of your quadratic equation given coefficients $\alpha$, $\beta$

This is actually a little trickier than it looks:
*  you can use numpy.sqrt(x) but be careful, it will fail if $x<0$ unless you call it as numpy.sqrt(x,dtype=complex)
*  The tests assume the two roots are sorted (hint: you can use numpy.sort())
*  After you have implemented your solution, you should remove the `raise NotImplementedError()` line (as you just implemented it)

<!-- grade_id: cell3621925556728684, points: 0, grade: false, solution: true -->
```python
def myroots(a, b):
    """Compute the two roots of the quadratic equation given by x^2 + ax + b = 0
       and return them as a sorted python list (or numpy array)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell5655123106444739, points: 5, grade: true, solution: false -->
```python
numpy.testing.assert_allclose(myroots(1.,0.), [-1.,0.])
numpy.testing.assert_allclose(myroots(1.,-2.), [-2.,1.])
numpy.testing.assert_allclose(myroots(-5.0,6.0), [2.,3.])
print('success: passed obvious tests')
```